Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Weekly mini-project 3, STATS102 Fall 2022 Session 1, Duke Kunshan University, Due at Sep 13th 1:15pm Beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

The data used in this assignment is courtesy of the FIFA World Cup Archive website and Kaggle.

Hint: pandas.DataFrame/Series.sort_values, pandas.Series.value_counts, pandas.DataFrame/Series.groupby pd.concat  pd.merge

(You may need to set parameters of functions for different conditions.)

# Problem 1 (0.5 point) 

Load `WorldCups.csv`, and return which country hosted the FIFA would cup with the highest attendance?  

<!-- grade_id: cell-aa92e1fe55529a21, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    df = pd.read_csv('WorldCups.csv')
    return df.iloc[df['Attendance'].str.replace('.','').astype(int).sort_values(ascending=False).head(1).index]['Country'].tolist()[0]
```

```python
import numpy as np
import pandas as pd
df=pd.read_csv('WorldCups.csv')
df.iloc[df['Attendance'].str.replace('.','').astype(int).sort_values(ascending=False).head(1).index]['Country'].tolist()[0]
```

Output:
```
'USA'
```

```python
import numpy as np
import pandas as pd
pd.read_csv('WorldCups.csv')
```

Output:
```
Year       Country      Winner      Runners-Up        Third  \
0   1930       Uruguay     Uruguay       Argentina          USA   
1   1934         Italy       Italy  Czechoslovakia      Germany   
2   1938        France       Italy         Hungary       Brazil   
3   1950        Brazil     Uruguay          Brazil       Sweden   
4   1954   Switzerland  Germany FR         Hungary      Austria   
5   1958        Sweden      Brazil          Sweden       France   
6   1962         Chile      Brazil  Czechoslovakia        Chile   
7   1966       England     England      Germany FR     Portugal   
8   1970        Mexico      Brazil           Italy   Germany FR   
9   1974       Germany  Germany FR     Netherlands       Poland   
10  1978     Argentina   Argentina     Netherlands       Brazil   
11  1982         Spain       Italy      Germany FR       Poland   
12  1986        Mexico   Argentina      Germany FR       France   
13  1990         Italy  Germany FR       Argentina        Italy   
```

<!-- grade_id: cell-151ed5ef3976c3f2, points: 0.5, grade: true, solution: false -->
```python
assert myfunction()=='USA'
```

# Problem 2 (0.5 point) 

Load `WorldCupMatches.csv`, and return the ratio of home team win.

You need to drop duplicate lines.

<!-- grade_id: cell-f87e606e70b0434b, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    df=pd.read_csv('WorldCupMatches.csv').dropna()
    length=len(df.index)
    win=0
    for i in range(len(df.index)):
        if df.iloc[i]['Home Team Goals']>df.iloc[i]['Away Team Goals']:
            win=win+1
        else:
            continue
    return win/length
```

```python
import numpy as np
import pandas as pd
df=pd.read_csv('WorldCupMatches.csv').dropna()
length=len(df.index)
win=0
for i in range(len(df.index)):
    if df.iloc[i]['Home Team Goals']>df.iloc[i]['Away Team Goals']:
        win=win+1
    else:
        continue
print(win/length)
```

Output:
```
0.571764705882353
```

```python
import numpy as np
import pandas as pd
df=pd.read_csv('WorldCupMatches.csv').dropna()
length=len(df.index)+2
type(df.iloc[0]['Home Team Goals'])
```

Output:
```
numpy.float64
```

```python
import numpy as np
import pandas as pd
pd.read_csv('WorldCupMatches.csv').dropna()
```

Output:
```
Year              Datetime                     Stage  \
0    1930.0  13 Jul 1930 - 15:00                    Group 1   
1    1930.0  13 Jul 1930 - 15:00                    Group 4   
2    1930.0  14 Jul 1930 - 12:45                    Group 2   
3    1930.0  14 Jul 1930 - 14:50                    Group 3   
4    1930.0  15 Jul 1930 - 16:00                    Group 1   
..      ...                   ...                       ...   
847  2014.0  05 Jul 2014 - 17:00             Quarter-finals   
848  2014.0  08 Jul 2014 - 17:00                Semi-finals   
849  2014.0  09 Jul 2014 - 17:00                Semi-finals   
850  2014.0  12 Jul 2014 - 17:00   Play-off for third place   
851  2014.0  13 Jul 2014 - 16:00                      Final   

                 Stadium             City Home Team Name  Home Team Goals  \
0                Pocitos      Montevideo          France              4.0   
1         Parque Central      Montevideo             USA              3.0   
2         Parque Ce
```

<!-- grade_id: cell-6c1d09ae5ed14ee6, points: 0.5, grade: true, solution: false -->
```python
assert abs(myfunction()-0.572)<0.01
```

# Problem 3 (0.5 point) 

Read `TerritoryReadMe.pdf` first to learn about the data in `PR.txt` and `TR.txt`.

Calulate the total number of occurrences of name $Aaliyah$ in Puerto Rico from 1998 to 2010.

<!-- grade_id: cell-5ffe97bd0fecf8d8, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def cnt_Aaliyah():
    # YOUR CODE HERE
    df=pd.read_table('PR.txt',sep=',',header=None)
    df.columns=['a','b','c','d','e']
    df=df[(df['d']=='Aaliyah')&(df['c']<=2010)&(df['c']>=1998)]
    return np.sum(df['e'])
```

```python
import numpy as np
import pandas as pd
df=pd.read_table('PR.txt',sep=',',header=None)
df
```

Output:
```
0  1     2         3    4
0      PR  F  1998     Paola  724
1      PR  F  1998   Genesis  500
2      PR  F  1998  Gabriela  447
3      PR  F  1998    Nicole  392
4      PR  F  1998   Alondra  344
...    .. ..   ...       ...  ...
21331  PR  M  2018    Yanuel    5
21332  PR  M  2018    Yaxiel    5
21333  PR  M  2018  Yendriel    5
21334  PR  M  2018      Yian    5
21335  PR  M  2018      Zael    5

[21336 rows x 5 columns]
```

```python
df.columns=['a','b','c','d','e']
df
```

Output:
```
a  b     c         d    e
0      PR  F  1998     Paola  724
1      PR  F  1998   Genesis  500
2      PR  F  1998  Gabriela  447
3      PR  F  1998    Nicole  392
4      PR  F  1998   Alondra  344
...    .. ..   ...       ...  ...
21331  PR  M  2018    Yanuel    5
21332  PR  M  2018    Yaxiel    5
21333  PR  M  2018  Yendriel    5
21334  PR  M  2018      Yian    5
21335  PR  M  2018      Zael    5

[21336 rows x 5 columns]
```

```python
df=df[(df['d']=='Aaliyah')&(df['c']<=2010)&(df['c']>=1998)]
df
```

Output:
```
a  b     c        d   e
609   PR  F  1998  Aaliyah   5
2264  PR  F  2001  Aaliyah  24
2898  PR  F  2002  Aaliyah  33
3554  PR  F  2003  Aaliyah  20
4170  PR  F  2004  Aaliyah  17
4899  PR  F  2005  Aaliyah   9
5377  PR  F  2006  Aaliyah  14
5944  PR  F  2007  Aaliyah  15
6575  PR  F  2008  Aaliyah  10
7072  PR  F  2009  Aaliyah  11
7685  PR  F  2010  Aaliyah   7
```

```python
np.sum(df['e'])
```

Output:
```
165
```

<!-- grade_id: cell-7715966e4288b04c, points: 0.5, grade: true, solution: false -->
```python
assert cnt_Aaliyah() == 165
```

# Problem 4 (0.5 point) 

Read `TerritoryReadMe.pdf` first to learn about the data in `PR.txt` and `TR.txt`.

In Puerto Rico and all other territories, find there are how many different names.

<!-- grade_id: cell-a88656ebc53cd614, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def cnt_names():
    # YOUR CODE HERE
    df1=pd.read_table('PR.txt',sep=',',header=None)
    df1.columns=['a','b','c','d','e']
    df2=pd.read_table('TR.txt',sep=',',header=None)
    df2.columns=['a','b','c','d','e']
    return len((df1['d'].value_counts())&(df2['d'].value_counts()))
```

<!-- grade_id: cell-2ef987bb4ebc53bf, points: 0.5, grade: true, solution: false -->
```python
assert cnt_names() == 3179
```

# Problem 5 (0.5 point) 

Read `TerritoryReadMe.pdf` first to learn about the data in `PR.txt` and `TR.txt`.

In Puerto Rico and all other territories, calculate the number of occurrences of the start letter of the name, return the top 3 letters in a list.

<!-- grade_id: cell-a08eb62f029129ba, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def find_letters():
    # YOUR CODE HERE
    headers = ['Territory','Gender','Year','Name','Number of occurrences']
    pr_df = pd.read_csv('PR.txt',header=None,names=headers)
    tr_df = pd.read_csv('TR.txt',header=None,names=headers)
    df = pd.concat([pr_df, tr_df])
    df['Start Letter'] = df['Name'].str[0]

    start_letter = df['Start Letter'].value_counts()
    letter_list = start_letter.sort_values(ascending=False)
    A=df['Start Letter'].value_counts().index[0:3]
    B=list(A)
    return B
```

```python
import numpy as np
import pandas as pd
headers = ['Territory','Gender','Year','Name','Number of occurrences']
pr_df = pd.read_csv('PR.txt',header=None,names=headers)
tr_df = pd.read_csv('TR.txt',header=None,names=headers)
df = pd.concat([pr_df, tr_df])
df['Start Letter'] = df['Name'].str[0]
    
start_letter = df['Start Letter'].value_counts()
letter_list = start_letter.sort_values(ascending=False)
```

<!-- grade_id: cell-f325580aaca09688, points: 0.5, grade: true, solution: false -->
```python
assert find_letters() == ['J', 'A', 'K']
```

# Problem 6 (0.5 point) 

Given the data of Chipotle Restaurant in `chipotle.tsv`, find the most-ordered item.

<!-- grade_id: cell-540b4be0ff74bcac, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def find_item():
    # YOUR CODE HERE
    df=pd.read_csv('chipotle.tsv',delimiter = '\t')
    return df.groupby('item_name')['quantity'].sum().sort_values(ascending=False).head(1).index.tolist()[0]
```

```python
import numpy as np
import pandas as pd
df=pd.read_csv('chipotle.tsv',delimiter = '\t')
length=len(df.groupby('item_name')['quantity'].sum())
df.groupby('item_name')['quantity'].sum().sort_values(ascending=False).head(1).index.tolist()[0]
```

Output:
```
'Chicken Bowl'
```

```python
import numpy as np
import pandas as pd
df=pd.read_csv('chipotle.tsv',delimiter = '\t')
df
```

Output:
```
order_id  quantity                              item_name  \
0            1         1           Chips and Fresh Tomato Salsa   
1            1         1                                   Izze   
2            1         1                       Nantucket Nectar   
3            1         1  Chips and Tomatillo-Green Chili Salsa   
4            2         2                           Chicken Bowl   
...        ...       ...                                    ...   
4617      1833         1                          Steak Burrito   
4618      1833         1                          Steak Burrito   
4619      1834         1                     Chicken Salad Bowl   
4620      1834         1                     Chicken Salad Bowl   
4621      1834         1                     Chicken Salad Bowl   

                                     choice_description item_price  
0                                                   NaN     $2.39   
1                                          [Clementine]     $3.
```

<!-- grade_id: cell-f8e3093630f197a9, points: 0.5, grade: true, solution: false -->
```python
assert find_item() == 'Chicken Bowl'
```