Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

quiz 8, STATS102 Fal 2021 Session 1, Duke Kunshan University, Sep 9th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

calculate and return the percentage of host country teams ('Country') that won the finals ('Winner') when there are greater or equal to 16 participating teams


\\( \frac{\text{how many times home teams that won the finals when there are greater or equal to 16 participating teams}}{\text{how many world cups when there are greater or equal to 16 participating teams}} \\)

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    cups = pd.read_csv('WorldCups.csv')
    B=np.sum(cups['QualifiedTeams']>=16)
    A=((cups['QualifiedTeams']>=16) & (cups['Country']==cups['Winner'])
    return A/B
```

Output:
```
File "<ipython-input-24-c63d51979b95>", line 7
    return A/B
         ^
SyntaxError: invalid syntax
```

<!-- grade_id: myfunction_test, points: 0.5, grade: true, solution: false -->
```python
assert abs(myfunction()-0.235)<0.01

```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-25-589473d7cf28> in <module>
----> 1 assert abs(myfunction()-0.235)<0.01

TypeError: unsupported operand type(s) for -: 'NoneType' and 'float'
```