Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

assignment 1, STATS102 Fall 2022 Session 1, Duke Kunshan University, due at Aug 25th 1:15PM beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

The term "mileage" often refers to distance per unit of fuel, such as miles per gallon (mpg) or litres per 100 kilometres (L/100 km). https://simple.wikipedia.org/wiki/Mileage

<!-- grade_id: calculate_mileage, points: 0, grade: false, solution: true -->
```python
def calculate_mileage(distance,fuel):
    
    """compute the mileage by dividing the travelled distance by the used fuel
       if the distance is negative or the fuel is not positive, then please raise a ValueError.    
    """   
   
    # YOUR CODE HERE
    #distance=input("distance?")
    #fuel=input("fuel?")
    Distance=float(distance)
    Fuel=float(fuel)
    if Distance<0 or Fuel <=0:
        raise ValueError
    else :
        return distance/fuel
```

<!-- grade_id: mileage_correct_input, points: 1, grade: true, solution: false -->
```python
"""Check that calculate_mileage function returns the correct output for several inputs (1 point)"""
assert calculate_mileage(100,5) == 20
assert calculate_mileage(50,4) == 12.5

```

<!-- grade_id: mileage_incorrect_input, points: 1, grade: true, solution: false -->
```python
"""Check that calculate_mileage function raises an error for invalid inputs (1 point)"""
try:
    calculate_mileage(-100,5)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
    
try:
    calculate_mileage(5,0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")

```