Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yins"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 8, STATS102 Fall 2022 Session 1, Duke Kunshan University, Due at Sep 7th 1:15pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

# Problem 1
Sort an 2D matrix by the nth column. For example:

Input A:
$$
 \begin{bmatrix}
   4 & 5 & 9 \\
   2 & 7 & 5 \\
   1 & 3 & 4 \\
 \end{bmatrix}
$$

after sorting A by the 2nd column, output:
$$
 \begin{bmatrix}
   1 & 3 & 4 \\
   4 & 5 & 9 \\
   2 & 7 & 5 \\
 \end{bmatrix}
$$

<!-- grade_id: cell-45ef485b9a06219e, points: 0, grade: false, solution: true -->
```python
import numpy as np

def myfunction(A, n):
    # YOUR CODE HERE
    A=A[A[:n].argsort()]
    return A
```

```python
A = np.arange(0,9)[::-1].reshape(3,3)
myfunction(A, 1)
```

Output:
```
array([[[2, 1, 0],
        [5, 4, 3],
        [8, 7, 6]]])
```

<!-- grade_id: cell-d5175c3e37339ebp, points: 1, grade: true, solution: false -->
```python
import numpy as np

A = np.arange(0,9)[::-1].reshape(3,3)
assert (myfunction(A, 1) == np.array([[2, 1, 0],
                                       [5, 4, 3],
                                       [8, 7, 6]])).all()
```

# Problem 2
Assuming the Estimated Tuition and Fees for the 2019-2020 Academic Year for one ECE master program is shown as follows.


please create and return the same DataFrame:

![](ece.png)

<!-- grade_id: cell-1b5420aa0da1a75f, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    df = pd.DataFrame(np.random.rand(3, 2),
                  index=pd.MultiIndex.from_arrays([[1,2,2], ['DKU','DKU','Duke']]),
                  columns=['International students','Domestic Students'])
    df.index.names=['year','where']
    df.iloc[0,0]='64500USD'
    df.iloc[1,0]='64500USD'
    df.iloc[2,0]='77250USD'
    df.iloc[0,1]='237000RMB'
    df.iloc[1,1]='237000RMB'
    df.iloc[2,1]= '77250USD'
    return df
```

<!-- grade_id: cell-b3c32ae0c1ec1fa8, points: 1, grade: true, solution: false -->
```python
#test 1 0.3pts
import numpy as np
import pandas as pd

tuition=myfunction()

assert len(tuition.index.levels)==2
assert len(tuition.index.names)==2
assert len(tuition.columns)==2
assert tuition['International students'][1]['DKU']=='64500USD'
assert tuition['Domestic Students'][2]['Duke']=='77250USD'
```