Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Weekly mini-project 1, STATS102 Fall 2022 Session 1, Duke Kunshan University, Due at Aug 29th 1:15pm Beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

# problem 1 (0.6 point) 
Given a list of integers, return indices of the two numbers such that they add up to a specific target.

Indices should be in an increasing order

You may assume that each input would have exactly one solution, and you may not use the same element twice.

<!-- grade_id: TwoSum, points: 0, grade: false, solution: true -->
```python
def TwoSum(nums, target):
    """
    :type nums: List[int]
    :type target: int
    :rtype: List[int]
    """
    # YOUR CODE HERE
    for i in range(len(nums)):
        for j in range(i+1, len(nums)):
            if nums[i] + nums[j] == target:
                return [i, j]
```

<!-- grade_id: TwoSum_test, points: 0.6, grade: true, solution: false -->
```python
'''this TwoSum function should pass the following tests'''
assert (TwoSum([2, 7, 11, 15],9)==[0,1])
assert (TwoSum([3,2,4],6)==[1,2])
assert (TwoSum([0,4,3,0],0)==[0,3])
assert (TwoSum([-1,-2,-3,-4,-5],-8)==[2,4])


```

# problem 2 (0.6 point) 
Given a 32-bit signed integer, reverse digits of an integer.
If the return number is outside this range \\( [-2^{31},2^{31}-1] \\), then return 0

example:  input 12345  output 54321

```python
a='4271y5r83928r204'
print(a[::-1])
```

Output:
```
402r82938r5y1724
```

<!-- grade_id: myreverse, points: 0, grade: false, solution: true -->
```python
def myreverse(x):
    ''':type x: int
       :rtype: int
    '''
    # YOUR CODE HERE
    z=str(x)
    if type(x)!= int:
        return 0
    elif -2**31>x or x>2**31:
        return 0
    elif z[0]=='-':
        z=z[1:]
        m=z[::-1]
        n=int(m)
        if -2**31>n or n>2**31:
            return 0
        else:
            return -n
    else:
        y=z[::-1]
        if y[0]=='0':
            o=int(y[1:])
            if -2**31>o or o>2**31:
                return 0
            else:
                return o
        else:
            p=int(y[:])
            if -2**31>p or p>2**31:
                return 0
            else:
                return p
```

<!-- grade_id: myreverse_test, points: 0.6, grade: true, solution: false -->
```python
assert(myreverse(120)==21)
assert(myreverse(123)==321)
assert(myreverse(-123)==-321)
assert(myreverse(1534236469)==0)
```

# Problem 3 (0.6 point) 
Write a function to find the longest common prefix string amongst a list of strings.

If there is no common prefix, return an empty string ''.

example: input ['what','where','who']     output 'wh'

<!-- grade_id: longestCommonPrefix, points: 0, grade: false, solution: true -->
```python
def longestCommonPrefix(strs):
    """
    :type strs: List[str]
    :rtype: str
    """
    # YOUR CODE HERE
    result = strs[0]
    for i, v in enumerate(strs[1:]):
        while not v.startswith(result):
            result = result[:-1]
            if not result:
                return ""
    return result
```

<!-- grade_id: longestCommonPrefix_test, points: 0.6, grade: true, solution: false -->
```python
assert longestCommonPrefix(["flower","flow","flight"])=='fl'
assert longestCommonPrefix(['what','where','who'] )=='wh'
```

# Problem 4 (0.6 point) 
Find the median of an integer list (don't directly call the system built-in function).

<!-- grade_id: cell-87d12d67d86866ab, points: 0, grade: false, solution: true -->
```python
def myMedian(int_list):
    """
    :type int_list: List[int]
    :rtype: float
    """
    # YOUR CODE HERE
    for i in range(len(int_list)):
        for j in range(len(int_list)-i-1):
            if int_list[j]>int_list[j+1]:
                int_list[j],int_list[j+1]=int_list[j+1],int_list[j]
    a=int_list
    if len(a)==1:
        return a[0]
    if len(a)%2==0:
        m=((a[int(len(a)/2-1)])+(a[int(len(a)/2)]))/2
        return m
    else:
        n=a[int((len(a)-1)/2)]
        return n
```

<!-- grade_id: cell-54d5a6421110bfda, points: 0.6, grade: true, solution: false -->
```python
assert myMedian([1,2,3,4]) == 2.5
assert myMedian([9,7,8]) == 8
assert myMedian([10]) == 10
assert myMedian([-1,2,-3,4]) == 0.5
```

# Problem 5 (0.6 point) 

Given a list of characters letters containing only lowercase letters and a target letter, write a function to find the smallest element in the list that is larger than the given target.

<!-- grade_id: cell-2bc93ebbf1d591fe, points: 0, grade: false, solution: true -->
```python
def findLetter(letters, target):
    """
    :type letters: List[str]
    :type target: str
    :rtype: str
    """
# YOUR CODE HERE
    m=[]
    for i in letters:
        if i > target:
            m.append(i)
        else:
            continue
    n=sorted(m)
    return n[0]
```

<!-- grade_id: cell-e00f009e6ec3b5c1, points: 0.6, grade: true, solution: false -->
```python
'''this findLetter function should pass the following tests'''
assert (findLetter(['a', 'b', 'c', 'd'],'a')=='b')
assert (findLetter(['g','i','j'],'h')=='i')
assert (findLetter(['c','c','p','o'],'c')=='o')
assert (findLetter(['d','c','f','k','g'],'a')=='c')
```

congratulations! Have a nice weekend!