Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Assignment 6, STATS102 Fall 2022 Session 1, Duke Kunshan University, Due at Sep 2nd 1:15pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

# Problem1 (0.5 point)

Given a string, find all possible positons of a substring in it.  

Hint: string.find function can take a second or even a third input parameters.

Docstring:
S.find(sub[, start[, end]]) -> int

Return the lowest index in S where substring sub is found,
such that sub is contained within S[start:end].  Optional
arguments start and end are interpreted as in slice notation.

Return -1 on failure.

<!-- grade_id: cell-f55edec9336761f9, points: 0, grade: false, solution: true -->
```python
def findSubstr(string, substr):
    
    # YOUR CODE HERE
    n=len(substr)
    alist=[]
    for i in range(len(string)+1):
        m=string.find(substr,i,i+n)
        if m==-1:
            continue
        else:
            alist.append(m)
    return alist
```

<!-- grade_id: cell-1eddbf23c0d24eb4, points: 0.5, grade: true, solution: false -->
```python
assert findSubstr('Hello', 'l')==[2, 3]
assert findSubstr('I love python', 'o')==[3, 11]
assert findSubstr('Banana', 'na')==[2, 4]
assert findSubstr('Good job', 'o')==[1, 2, 6]
```

# Problem 2 (0.5 point)

Write a program to encode the string by spliting it by the middle and reverse the two substrings.

<!-- grade_id: cell-6e14d84ed557fa4d, points: 0, grade: false, solution: true -->
```python
def encodeString(ori_str):
    # YOUR CODE HERE
    try:
        n=len(ori_str)
        left=ori_str[:n//2]
        right=ori_str[n//2:]
        Left=left[::-1]
        Right=right[::-1]
        return Left+Right
    except:
        n=len(ori_str)
        n%2!=0
        pass
```

<!-- grade_id: cell-4ccb27702e3e4b9d, points: 0.5, grade: true, solution: false -->
```python
assert encodeString('AaAaBbBb') == 'aAaAbBbB'
assert encodeString('This lines must be printed backwards and splitted in the middle.') \
    == 'wkcab detnirp eb tsum senil sihT.elddim eht ni dettilps dna sdra'
assert encodeString(encodeString('This lines must be printed backwards and splitted in the middle.')) \
    == 'This lines must be printed backwards and splitted in the middle.'
assert encodeString('This is a very secret message!') \
    == ' yrev a si sihT!egassem terces'
assert encodeString(encodeString('This is a very secret message!')) \
    == 'This is a very secret message!'
assert encodeString('Happy programming~') \
    == 'orp yppaH~gnimmarg'
assert encodeString(encodeString('Happy programming~')) \
    == 'Happy programming~'
```

# Problem 3 (1 point)

Create a $9*9$ matrix:
$$
 \begin{bmatrix}
   1 & 0 & 3 & 0 & 5 & 0 & 7 & 0 & 9\\
   0 & 2 & 0 & 4 & 0 & 6 & 0 & 8 & 0\\
   11 & 0 & 13 & 0 & 15 & 0 & 17 & 0 & 19\\
   0 & 10 & 0 & 12 & 0 & 14 & 0 & 16 & 0\\
   21 & 0 & 23 & 0 & 25 & 0 & 27 & 0 & 29\\
   0 & 18 & 0 & 20 & 0 & 22 & 0 & 24 & 0\\
   31 & 0 & 33 & 0 & 35 & 0 & 37 & 0 & 39\\
   0 & 26 & 0 & 28 & 0 & 30 & 0 & 32 & 0\\
   41 & 0 & 43 & 0 & 45 & 0 & 47 & 0 & 49\
  \end{bmatrix}
$$

write a function to output this numpy matrix

<!-- grade_id: cell-dd0f5a15a60d8cf2, points: 0, grade: false, solution: true -->
```python
import numpy as np
def createMatrix():
    # YOUR CODE HERE
    matrix=np.full((9,9),0)
    matrix[0,0]=1
    matrix[0,2]=3
    matrix[0,4]=5
    matrix[0,6]=7
    matrix[0,8]=9
    for m in range(2,9,2):
        matrix[m,0]=matrix[m-2,0]+10
        matrix[m,2]=matrix[m-2,2]+10
        matrix[m,4]=matrix[m-2,4]+10
        matrix[m,6]=matrix[m-2,6]+10
        matrix[m,8]=matrix[m-2,8]+10
    matrix[1,1]=2
    matrix[1,3]=4
    matrix[1,5]=6
    matrix[1,7]=8
    for n in range(3,9,2):
        matrix[n,1]=matrix[n-2,1]+8
        matrix[n,3]=matrix[n-2,3]+8
        matrix[n,5]=matrix[n-2,5]+8
        matrix[n,7]=matrix[n-2,7]+8
    return matrix
```

<!-- grade_id: cell-8ad47c5cae4358fb, points: 1, grade: true, solution: false -->
```python
import numpy as np

assert (createMatrix() == np.array([[ 1,  0,  3,  0,  5,  0,  7,  0,  9],
                                       [ 0,  2,  0,  4,  0,  6,  0,  8,  0],
                                       [11,  0, 13,  0, 15,  0, 17,  0, 19],
                                       [ 0, 10,  0, 12,  0, 14,  0, 16,  0],
                                       [21,  0, 23,  0, 25,  0, 27,  0, 29],
                                       [ 0, 18,  0, 20,  0, 22,  0, 24,  0],
                                       [31,  0, 33,  0, 35,  0, 37,  0, 39],
                                       [ 0, 26,  0, 28,  0, 30,  0, 32,  0],
                                       [41,  0, 43,  0, 45,  0, 47,  0, 49]])).all()
```