Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

quiz 9, STATS102 Fall 2022 Session 1, Duke Kunshan University, Sep 7th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Please create and return one DataFrame shown as follows:

![](Quiz11.png)

Hint: pd.MultiIndex

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
#option 0.
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    df = pd.DataFrame(np.random.rand(4,1),
              index=pd.MultiIndex.from_arrays([[1998,1998,2018,2018], ['Paola','Genesis','Valentina','Victoria']]),
              columns=['Number of occurrences'])
    df.index.names=['Year','Name']
    df.iloc[0,0]=724
    df.iloc[1,0]=500
    df.iloc[2,0]=241
    df.iloc[3,0]=205
    return df
```

```python
#option 1.
import numpy as np
import pandas as pd
def myfunction():
    index=pd.MultiIndex.from_arrays([[1998,1998,2018,2018], ['Paola','Genesis','Valentina','Victoria']],names=['Year','Name'])
    data=np.array([724,500,241,205])
    name_cnt=pd.DataFrame(data, index=myindex, columns=['Numbers of occurrences'])
    return name_cnt
```

```python
import numpy as np
import pandas as pd
df = pd.DataFrame(np.random.rand(4,1),
              index=pd.MultiIndex.from_arrays([[1998,1998,2018,2018], ['Paola','Genesis','Valentina','Victoria']]),
              columns=['Number of occurrences'])
df.index.names=['Year','Name']
df.iloc[0,0]='724'
df.iloc[1,0]='500'
df.iloc[2,0]='241'
df.iloc[3,0]='205'
df
```

Output:
```
Number of occurrences
Year Name                           
1998 Paola                       724
     Genesis                     500
2018 Valentina                   241
     Victoria                    205
```

```python
#option 3.
import numpy as np
import pandas as pd
year=pd.Series([1998,1998,2018,2018])
name=pd.Series(['Paola','Genesis','Valentina','Victoria'])
number=pd.Series([724,500,241,205])
df=pd.DataFrame({'Year':year,'Name':name,'Number of occurrences':number})
df=df.index.names=['Year','Name']
df
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-9-6c046d30acef> in <module>
      6 number=pd.Series([724,500,241,205])
      7 df=pd.DataFrame({'Year':year,'Name':name,'Number of occurrences':number})
----> 8 df=df.index.names=['Year','Name']
      9 df

AttributeError: 'builtin_function_or_method' object has no attribute 'names'
```

<!-- grade_id: myfunction_test1, points: 0.5, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd

name_cnt = myfunction()
assert len(name_cnt.index.levels)==2
assert len(name_cnt.index.names)==2
assert len(name_cnt.columns)==1
assert (name_cnt['Number of occurrences'] == [724,500,241,205]).all()
```

Output:
```
<ipython-input-2-e061ffc305b4> in <module>
      2 import pandas as pd
      3 
----> 4 name_cnt = myfunction()
      5 assert len(name_cnt.index.levels)==2
      6 assert len(name_cnt.index.names)==2

<ipython-input-1-f2428db6c27c> in myfunction()
      5     index=pd.MultiIndex.from_arrays([[1998,1998,2018,2018], ['Paola','Genesis','Valentina','Victoria']],names=['Year','Name'])
      6     data=np.array([724,500,241,205])
----> 7     name_cnt=pd.DataFrame(data, index=myindex, columns=['Numbers of occurrences'])
      8     return name_cnt

NameError: name 'myindex' is not defined
```