Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Assignment 11, STATS102 Fall 2022 Session 1, Duke Kunshan University, Due at Sep 14th 1:15pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

# Problem 1

Load <b>WorldCupMatches.csv</b>, calculate and return a DataFrame representing the top 5 teams descendingly ordered by the number of goals in a single world cup.

![](assignment13.png)

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    
```

```python
import numpy as np
import pandas as pd
pd.read_csv('WorldCupMatches.csv').dropna()
```

```python
import numpy as np
import pandas as pd
df=pd.read_csv('WorldCupMatches.csv').dropna().drop_duplicates()
df1=df.loc[:,['Year','Home Team Name','Home Team Goals']]
df1.columns=['Year','Team Name','Goals in a single cup']
df2=df.loc[:,['Year','Away Team Name','Away Team Goals']]
df2.columns=['Year','Team Name','Goals in a single cup']
df=pd.concat([df1,df2], ignore_index=True)
df=df.groupby(by=['Year','Team Name']).sum()
df.sort_values(by=['Goals in a single cup'],ascending=False).head(5)
```

```python
import numpy as np
import pandas as pd
df=pd.read_csv('WorldCupMatches.csv').dropna()
df1=df.loc[:,['Year','Home Team Name','Home Team Goals']]
df1.columns=['Year','Team Name','Goals in a single cup']
df2=df.loc[:,['Year','Away Team Name','Away Team Goals']]
df2.columns=['Year','Team Name','Goals in a single cup']
display(df1,df2)
```

```python
df=pd.concat([df1,df2], ignore_index=True)
df
```

```python
df=df.groupby(by=['Year','Team Name']).sum()
df
```

```python
df.sort_values(by=['Goals in a single cup'],ascending=False).head(5)

```

<!-- grade_id: cell-044b32eea7f84461, points: 1.5, grade: true, solution: false -->
```python
res=myfunction()
assert res.shape==(5,1)
assert (res.values==np.array([[27.],[25.],[23.],[22.],[19.]])).all()
assert int(res.loc[1954,'Hungary'])==27
assert int(res.loc[1954,'Germany FR'])==25
assert int(res.loc[1958,'France'])==23
assert int(res.loc[1950,'Brazil'])==22
assert int(res.loc[1970,'Brazil'])==19
res
```

# Problem 2

Load <b>artists.csv</b>, calculate and return how many artists's lifespans are 30 to 60 years (excluding 30).

Data reference: https://www.kaggle.com/ikarus777/best-artworks-of-all-time

Hint:
![](assignment13_2.png)

<!-- grade_id: cell-598931e3074a3182, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    df=pd.read_csv('artists.csv')
    born=df['years'].str.extract('(\d+)\s.*').astype(int)
    died=df['years'].str.extract('\s(\d+).*$').astype(int)
    lifespan=died-born
    return np.sum((lifespan>30)&(lifespan<=60))[0]
```

```python
import numpy as np
import pandas as pd
df=pd.read_csv('artists.csv')
born=df['years'].str.extract('(\d+)\s.*').astype(int)
died=df['years'].str.extract('\s(\d+).*$').astype(int)
lifespan=died-born
np.sum((lifespan>30)&(lifespan<=60))[0]
```

Output:
```
19
```

```python
import numpy as np
import pandas as pd
df=pd.read_csv('artists.csv')
born=df['years'].str.extract('(\d+)\s.*').astype(int)
died=df['years'].str.extract('\s(\d+).*$').astype(int)
lifespan=died-born
lifespan.columns=['lifespan']
period=pd.cut(lifespan['lifespan'], [30,60,100])
period=period.value_counts()
period=period.to_frame()
period.columns=['numbers']
period
```

Output:
```
numbers
(60, 100]       31
(30, 60]        19
```

```python
period.loc('(30, 60]')
```

Output:
```
/home/dku-dasc-01/anaconda3/lib/python3.6/site-packages/pandas/core/indexing.py in __call__(self, axis)
    110 
    111         if axis is not None:
--> 112             axis = self.obj._get_axis_number(axis)
    113         new_self.axis = axis
    114         return new_self

/home/dku-dasc-01/anaconda3/lib/python3.6/site-packages/pandas/core/generic.py in _get_axis_number(cls, axis)
    409             except KeyError:
    410                 pass
--> 411         raise ValueError("No axis named {0} for object type {1}".format(axis, cls))
    412 
    413     @classmethod

ValueError: No axis named (30, 60] for object type <class 'pandas.core.frame.DataFrame'>
```

```python
period.loc['(60, 100]','numbers']
```

Output:
```
/home/dku-dasc-01/anaconda3/lib/python3.6/site-packages/pandas/core/indexes/category.py in get_loc(self, key, method)
    517         array([False,  True, False,  True], dtype=bool)
    518         """
--> 519         code = self.categories.get_loc(key)
    520         code = self.codes.dtype.type(code)
    521         try:

/home/dku-dasc-01/anaconda3/lib/python3.6/site-packages/pandas/core/indexes/interval.py in get_loc(self, key, method)
    864             except TypeError:
    865                 # scalar is not comparable to II subtype --> invalid label
--> 866                 raise KeyError(key)
    867 
    868         matches = mask.sum()

KeyError: '(60, 100]'
```

<!-- grade_id: cell-e10c724844514f74, points: 0.5, grade: true, solution: false -->
```python
assert myfunction()==19
```