Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Lab 3, STATS102 Fall2022 Session 1, Duke Kunshan University, Sep 8th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

# Problem 1

Loading `'CAvideos.csv`, find the video with the most views and return its title.


<!-- grade_id: strictly_ascending_list, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def most_views():
    # YOUR CODE HERE
    df=pd.read_csv('CAvideos.csv')
    Df=df.sort_values(by=['views'],ascending=False).head(1)
    return Df.iloc[[0],[2]].values
```

<!-- grade_id: strictly_ascending_list1, points: 0, grade: true, solution: false -->
```python
assert most_views() == 'Ed Sheeran - Perfect (Official Music Video)'
```

# Problem 2

Loading `'CAvideos.csv`, find the top 3 video with the most comments and return their titles in a list.

<!-- grade_id: cell-e0009896e4596007, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def top3_comments():
    # YOUR CODE HERE
    df=pd.read_csv('CAvideos.csv')
    Df=df.sort_values(by=['comment_count'],ascending=False).head(3)
    return Df['title'].tolist()
```

<!-- grade_id: cell-7debb17a223463d1, points: 0, grade: true, solution: false -->
```python
assert top3_comments() == ['Eminem - Walk On Water (Audio) ft. Beyoncé',
                             'Ed Sheeran - Perfect (Official Music Video)',
                             'Jake Paul - Saturday Night (Song) feat. Nick Crompton & Chad Tepper (Official Music Video)']
```

# Problem 3

Loading `'CAvideos.csv`, find the tag with the highest frequency.

All tags should be lowercase.

<!-- grade_id: cell-be2566487c305f77, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def highest_freq_tag():
    # YOUR CODE HERE
    df=pd.read_csv('CAvideos.csv')
    df['tags'].str.replace('\','')
```

Output:
```
File "<ipython-input-76-e0cc5104a4cd>", line 7
    df['tags'].str.replace('\','')
                                  ^
SyntaxError: EOL while scanning string literal
```

<!-- grade_id: cell-be2566487c305f77, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def highest_freq_tag():
    # YOUR CODE HERE
    df=pd.read_csv('CAvideos.csv')
    np.concatenate(df['tags'].str.replace('\','').str.split('|'))
```

Output:
```
File "<ipython-input-76-e0cc5104a4cd>", line 7
    df['tags'].str.replace('\','')
                                  ^
SyntaxError: EOL while scanning string literal
```

```python
import numpy as np
import pandas as pd
from collections import Counter
Df=pd.read_csv('CAvideos.csv')
df=Df['tags'].str.lower()
list=[]
for i in range(199):
    list=list+df[i].split('|')
tag=[]
dict={}
for j in list:
    if j in tag:
        dict[j]=dict.get(j,0)+1
    else:
        tag=tag+[j]
        dict[j]=dict.get(j,0)+1
max(dict.values())
for k,v in dict.items():
    if v==max(dict.values()):
        print(str(k))
```

Output:
```
"comedy"
```

<!-- grade_id: cell-f8b6f3af9b05ae45, points: 0, grade: true, solution: false -->
```python
assert highest_freq_tag() == 'comedy'
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-68-15849fad40e9> in <module>
----> 1 assert highest_freq_tag() == 'comedy'

AssertionError:
```

# Problem 4

Loading `'CAvideos.csv`, find there are how many videos of which the like ratio is more than $0.9$ and the number of likes is more than $8000$. (Ignore `Nan`) 

$$like\_ratio =  \frac{\text{the number of likes}}{\text{the number of likes + the number of dislikes}} $$

<!-- grade_id: cell-2dd857483e104c7d, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def analyze_likes():
    # YOUR CODE HERE
    df=pd.read_csv('CAvideos.csv')
    return (df['likes']/(df['likes']+df['dislikes'])>0.9)&(df['likes']>8000)
```

<!-- grade_id: cell-751309db16b03a89, points: 0, grade: true, solution: false -->
```python
assert analyze_likes() == 96
```

Output:
```
<ipython-input-82-bb0c373c7b3c> in <module>
----> 1 assert analyze_likes() == 96

/home/dku-dasc-01/anaconda3/lib/python3.6/site-packages/pandas/core/generic.py in __nonzero__(self)
   1553             "The truth value of a {0} is ambiguous. "
   1554             "Use a.empty, a.bool(), a.item(), a.any() or a.all().".format(
-> 1555                 self.__class__.__name__
   1556             )
   1557         )

ValueError: The truth value of a Series is ambiguous. Use a.empty, a.bool(), a.item(), a.any() or a.all().
```