Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Assignment 13, STATS102 Fall2022 Session 1, Duke Kunshan University, Due Sep 16th 1:15pm
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

load WorldCupMatches.csv and WorldCupPlayers.csv to calculate all the matches that have goal in the first minute, and output the following dataframe.

![](ex5.png)

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    df1=pd.read_csv('WorldCupMatches.csv').set_index('MatchID').dropna(how='all').drop_duplicates()
    df2=pd.read_csv('WorldCupPlayers.csv').set_index('MatchID').dropna(how='all').drop_duplicates()

    s=df2['Event'].dropna()
    a=[]
    alist=[]
    for i in range(len(s)):
        if "G1'" in s.iloc[i]:
            a.append(i)
    for m in a:
        alist.append(s.index[m])
    df=df2.loc[alist][['Player Name','Event']].dropna(how='any')
    num=[]
    for m in range(len(df)):
        if df['Event'].values[m].split('\'')[0]=='G1':
            num.append(m)
    DF=df.iloc[num]
    chart=pd.merge(df1,DF,on='MatchID')
    chart=chart[['Year','Datetime','Home Team Name','Home Team Goals','Away Team Goals','Away Team Name','Player Name']]
    return chart
```

```python
myfunction()
```

Output:
```
Year              Datetime  Home Team Name  Home Team Goals  \
MatchID                                                                      
1098.0       1930.0  14 Jul 1930 - 14:50          Romania              3.0   
1105.0       1934.0  07 Jun 1934 - 18:00          Germany              3.0   
1146.0       1938.0  05 Jun 1938 - 17:00           France              3.0   
1176.0       1938.0  16 Jun 1938 - 18:00          Hungary              5.0   
1470.0       1962.0  03 Jun 1962 - 15:00          Hungary              6.0   
1544.0       1962.0  07 Jun 1962 - 15:00           Mexico              3.0   
1702.0       1966.0  23 Jul 1966 - 15:00         Portugal              5.0   
2347.0       1978.0  02 Jun 1978 - 13:45            Italy              2.0   
878.0        1982.0  16 Jun 1982 - 17:15          England              3.0   
8763.0       1998.0  24 Jun 1998 - 21:00          Nigeria              1.0   
43950063.0   2002.0  29 Jun 2002 - 20:00   Korea Republic              2.0   
3
```

<!-- grade_id: myfunction_test1, points: 2, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
res=myfunction()
assert res.shape==(12,7)
assert res['Home Team Goals'].sum()==37
assert res['Away Team Goals'].sum()==20
```