Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Mid-term, STATS102 Fall Session 1, Duke Kunshan University, Feb 19th 8pm-10pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

# for autograding questions, there is no partial grades if the auto-testing (assertion) is failed. TA will manually look at each solution, so please do not abuse the auto-grading system, e.g. return ['what', 'who', 'when'] for Problem 4

# Problem 1: (12 points)
Implement your own function to find the median of an integer list. (you need to implement the median function by yourself, using existing median function is not allowed)


The median is the middle number in a sorted, ascending or descending, list of numbers and can be more descriptive of that data set than the average. If there is an odd amount of numbers, the median value is the number that is in the middle, with the same amount of numbers below and above.
If there is an even amount of numbers in the list, the middle pair must be determined, added together, and divided by two to find the median value.

<!-- grade_id: myMedian, points: 0, grade: false, solution: true -->
```python
def myMedian(int_list):
    """
    :input  type int_list: List[int]
    :return type: float
    """
    # YOUR CODE HERE
    if len(int_list)%2==0:
        a=sorted(int_list)
        half_length=len(int_list)//2
        return (a[half_length-1]+a[half_length])/2
    else:
        a=sorted(int_list)
        return a[len(int_list)//2]
```

<!-- grade_id: myMedian_test, points: 12, grade: true, solution: false -->
```python
assert myMedian([1,2,3,4]) == 2.5
assert myMedian([9,7,8]) == 8
assert myMedian([10]) == 10
assert myMedian([-1,2,-3,4]) == 0.5
#no hidden test here
```

# Problem 2:  (13 points)
Given a list of characters letters containing only lowercase letters and a target letter, write a function to find the smallest element in the list that is larger than the given target. if there is no element in the list that is larger than the given target, return None


<!-- grade_id: cell-5ed38e5f0a8a8433, points: 0, grade: false, solution: true -->
```python
def findLetter(letters, target):
    """
    :type letters: List[str]
    :type target: str
    :rtype: str or None
    """
    # YOUR CODE HERE
    m=[]
    for i in letters:
        if i > target:
            m.append(i)
        else:
            continue
    n=sorted(m)
    if len(n)==0:
        return None
    else:
        return n[0]
```

<!-- grade_id: cell-943322dfb58d72d3, points: 13, grade: true, solution: false -->
```python
'''this findLetter function should pass the following tests'''
assert (findLetter(['a', 'b', 'c', 'd'],'a')=='b')
assert (findLetter(['g','i','h'],'j')==None)
assert (findLetter(['c','c','p','o'],'c')=='o')
assert (findLetter(['d','c','f','k','g'],'a')=='c')
assert(findLetter( ['c', 'e', 'd'],'g')==None)
#no hidden test here
```

# Problem 3:  (15 points)
Write a program to calculate the square of Euclidean distance between two RGB colors:

$$distance^2=(R_2-R_1)^2 + (G_2-G_1)^2 + (B_2-B_1)^2$$
  
where both R, G and B values are integer numbers between 0-255 (closed interval [0,255]).

Make sure it handles the case where R, G and B values are not in 0-255 by raising a `ValueError`.

<!-- grade_id: cell-d1b57c91046ae247, points: 0, grade: false, solution: true -->
```python
def distance_square(R1, G1, B1, R2, G2, B2):
    """Calculate the square of Euclidean distance between two RGB colors.
    """   
    # YOUR CODE HERE
    if (R1<0 or R1>255):
        raise ValueError
    elif (R2<0 or R2>255):
        raise ValueError
    elif (G1<0 or G1 >255):
        raise ValueError
    elif (G2<0 or G2 >255):
        raise ValueError
    elif (B1<0 or B1 >255):
        raise ValueError
    elif (B2<0 or B2 >255):
        raise ValueError
    else:
        return (R2-R1)**2+(G2-G1)**2+(B2-B1)**2
```

<!-- grade_id: cell-c9ec8ff72d658591, points: 5, grade: true, solution: false -->
```python
"""Check that distance returns the correct output for several inputs"""
# for valid testing intputs, 5 points
assert distance_square(0,128,0,255,0,0) == 81409
assert distance_square(128,0,128,0,0,255) == 32513
#no hidden test here
```

<!-- grade_id: cell-564c152a3b0c8f3b, points: 10, grade: true, solution: false -->
```python
"""Check that distance raises an error for invalid inputs"""
# for invalid testing intputs, 10 points
try:
    distance_square(-1,1,2,3,4,6)
except ValueError:
    pass
else:
    raise AssertionError("Error input")

try:
    distance_square(0,1,2,3,4,256)
except ValueError:
    pass
else:
    raise AssertionError("Error input")
#no hidden test here
```

# Problem 4: (15 points)

In the file `questions.txt`, there are 9095 questions from the Stanford Question Answering Dataset, whose question words are in the list `['what', 'who', 'when', 'where', 'how', 'how_many', 'how_much', 'how_old']`. 

Now, please write a program to calculate the frequency of each question word and output a list containing the top 3 question words with high frequency in the occurance decending order.

Notice: For question like "How old are you?", we consider its question word is "how_old" rather than "how".

<!-- grade_id: cell-b413f6f29108a8f0, points: 0, grade: false, solution: true -->
```python
def sort_list_question_words():
    '''output a list of strs'''
    question_words = ['what', 'who', 'when', 'where', 'how', 'how_many', 'how_much', 'how_old']
    handle = open('questions.txt', 'r', encoding='UTF-8')
    #for line in handle:

    # YOUR CODE HERE
    words=[]
    for line in handle:
        if line.startswith('What') :
            words.append('what')
        if line.startswith('Who') :
            words.append('who')
        if line.startswith('When') :
            words.append('when')
        if line.startswith('Where') :
            words.append('where')
        if line.startswith('How') :
            t=line.split(' ')
            if t[1]=='many':
                words.append('how_many')
            elif t[1]=='much':
                words.append('how_much')
            elif t[1]=='old':
                words.append('how_old')
            else:
                words.append('how')
    counts = dict()
    question_words = ['what', 'who', 'when', 'where', 'how', 'how_many', 'how_much', 'how_old']
    for word in words :
        if word not in counts: 
            counts[word] = 1
        else :
            counts[word] = counts[word] + 1
    #print(counts)
    #counts.items()
    alist=[]
    for i in [0,1,2]:
        alist.append(sorted(counts.items(),key=lambda x: x[1])[i][0])
    return alist
```

```python
'''output a list of strs'''
question_words = ['what', 'who', 'when', 'where', 'how', 'how_many', 'how_much', 'how_old']
handle = open('questions.txt', 'r', encoding='UTF-8')
#for line in handle:

# YOUR CODE HERE
words=[]
for line in handle:
    if line.startswith('What') :
        words.append('what')
    if line.startswith('Who') :
        words.append('who')
    if line.startswith('When') :
        words.append('when')
    if line.startswith('Where') :
        words.append('where')
    if line.startswith('How') :
        t=line.split(' ')
        if t[1]=='many':
            words.append('how_many')
        elif t[1]=='much':
            words.append('how_much')
        elif t[1]=='old':
            words.append('how_old')
        else:
            words.append('how')
counts = dict()
question_words = ['what', 'who', 'when', 'where', 'how', 'how_many', 'how_much', 'how_old']
for word in words :
    if word not in counts: 
        counts[word] = 1
    else :
        counts[word] = counts[word] + 1
#print(counts)
#counts.items()
alist=[]
for i in [0,1,2]:
    alist.append(sorted(counts.items(),key=lambda x: x[1])[i][0])
print (alist)
```

Output:
```
['how_old', 'how_much', 'where']
```

```python
assert sort_list_question_words()==['what', 'who', 'when']
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-6-97c1a58f618b> in <module>
----> 1 assert sort_list_question_words()==['what', 'who', 'when']

AssertionError:
```

<!-- grade_id: cell-4610db83f2fa2344, points: 15, grade: true, solution: false -->
```python
#hidden test here (you need to return a python list with three strings, the answer is not shown to you)

```

# Problem 5: (10 points)

For an input scalar variable x and positive integer N, use numpy to calculate and return \\( \Sigma_{n=0}^{N}\frac{x^{n}}{n!}\\), we will compare the results with \\(e^{x}\\) according to different N. 



<!-- grade_id: cell-b847b4cb8448fda8, points: 0, grade: false, solution: true -->
```python
import numpy as np

def myfunction(x, N):
    # YOUR CODE HERE
    k=np.arange(0,N+1)
    n=0
    for k in range(0,N+1):
        m=x**k/np.math.factorial(k)
        n=n+m
    return n
```

<!-- grade_id: cell-1e2cf9a87af56dc9, points: 10, grade: true, solution: false -->
```python
import numpy as np

assert np.abs(myfunction(1,10)-np.exp(1))<0.02*np.exp(1)
assert np.abs(myfunction(0.5,10) -np.exp(0.5))<0.02*np.exp(0.5)
assert np.abs(myfunction(5,20)-np.exp(5))<0.02*np.exp(5)
#no hidden test here
```

# Problem 6: (15 points)
Find the product of the maximum and minimum integer numbers occuring in `regex_sum_320982.txt` and `regex_sum_42.txt`.

the integer number may be negative


<!-- grade_id: cell-c49a97195af86cda, points: 0, grade: false, solution: true -->
```python
import re

def find_product(file):
    # YOUR CODE HERE
   # if file=='regex_sum_320982.txt':
    fhand1=open(file, 'r', encoding='UTF-8')
    list1=[]
    for line in fhand1:
        number=re.findall('([-a]*\d+)',line)
        if number==[]:
            continue
        else:
            list1=list1+number
    list1=list(map(int,list1))
    list1=sorted(list1)
    return list1[0]*list1[::-1][0]
    
  #  elif file=='regex_sum_42.txt':
#        fhand2=open('regex_sum_42.txt', 'r', encoding='UTF-8')
#        list2=[]
#        for line in fhand2:
#            number=re.findall('([-a]*\d+)',line)
#            if number==[]:
#                continue
#            else:
#                list2=list2+number
#        list2=list(map(int,list2))
#        list2=sorted(list2)
#    return list2[0]*list2[::-1][0]
```

```python
assert find_product('regex_sum_42.txt') == -313440
assert find_product('regex_sum_320982.txt') == -415548

```

<!-- grade_id: cell-e2316df645b4435b, points: 15, grade: true, solution: false -->
```python
#hidden test here

```

# Problem 7: (10 points)
given one dataframe as follows:

full_monte = pd.DataFrame(data=['Graham Chapman', 'John Cleese', 'Terry Gilliam'],columns=['Name'])

![](72.png)

please return one new dataframe that has an additioal column of last name, calculated from the Name column as follows

![](71.png)

<!-- grade_id: cell-18d02e30d1f76a70, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    full_monte = pd.DataFrame(data=['Graham Chapman', 'John Cleese', 'Terry Gilliam'],columns=['Name'])
    # YOUR CODE HERE
    alist=[]
    for i in range(len(full_monte)):
        last=full_monte["Name"][i].split(' ')[::-1][0]
        alist.append(last)
    full_monte["Last Name"]=alist
    return full_monte
```

<!-- grade_id: cell-14f5cfc9c7cdd63e, points: 10, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
full_monte2=myfunction()
assert full_monte2['Last Name'].iloc[0]=='Chapman'
assert full_monte2['Last Name'].iloc[1]=='Cleese'
assert full_monte2['Last Name'].iloc[2]=='Gilliam'
#no hidden test here
```

# Problem 8 (10 point)

Write a program to determine which day of the year. 

Input: the date of the year, with the format `year-month-day`, e.g. `2018-9-4`.

Output: which day of the year.



Hint:
2 criteria must be taken into account to identify leap years:
1. The year is evenly divisible by 4, and can NOT be evenly divided by 100;
2. The year is evenly divisible by 400.

<!-- grade_id: cell-d786e1ebf0fe1058, points: 0, grade: false, solution: true -->
```python
def dayOfYear(date):
    # YOUR CODE HERE
    year=int(re.findall('(\d+)[a-]',date)[0])
    month=int(re.findall('[a-](\d+)[a-]',date)[0])
    day=int(re.findall('[a-](\d+)',date)[-1])
    if (year%4==0 and year%100!=0):
        if month>2:
            if month%2!=0:
                days=(month-1)//2*31+(month-1)//2*30-1+day
            else:
                days=month//2*31+(month//2-1)*30-1+day
        if month==2:
            days=31+day
        if month==1:
            days=day
    elif (year%400==0):
        if month>2:
            if month%2!=0:
                days=(month-1)//2*31+(month-1)//2*30-1+day
            else:
                days=month//2*31+(month//2-1)*30-1+day
        if month==2:
            days=31+day
        if month==1:
            days=day
    else:
        if month>2:
            if month%2!=0:
                days=(month-1)//2*31+(month-1)//2*30-2+day
            else:
                days=month//2*31+(month//2-1)*30-2+day
        if month==2:
            days=31+day
        if month==1:
            days=day
    return days
```

<!-- grade_id: cell-fa2968e3acd0be8d, points: 5, grade: true, solution: false -->
```python
assert dayOfYear('2012-10-1') == 275
assert dayOfYear('2000-1-1') == 1
assert dayOfYear('1900-2-28') == 59
assert dayOfYear('1900-3-1') == 60
assert dayOfYear('1904-3-1') == 61
assert dayOfYear('2012-10-11') == 285
```

<!-- grade_id: cell-0f40e390ff2f85ad, points: 5, grade: true, solution: false -->
```python
try:
    dayOfYear('year-month-day') # error input formate
    dayOfYear('2019_8_9') # error input format
    dayOfYear('2019-8') # incompleted input
    dayOfYear('2016-2-30') # no 30 days in the Feb. of a leap year
    dayOfYear('2018-2-29') # no 29 days in the Feb. of a normal year
    dayOfYear('2016-22-30') # no 22 months in a year
    dayOfYear('2016-10-36') # no 36 days in a month
except ValueError:
    pass
else:
    raise AssertionError("Error input")
```

Output:
```
<ipython-input-197-97b5be46dfbc> in <module>
      1 try:
----> 2     dayOfYear('year-month-day') # error input formate
      3     dayOfYear('2019_8_9') # error input format
      4     dayOfYear('2019-8') # incompleted input
      5     dayOfYear('2016-2-30') # no 30 days in the Feb. of a leap year

<ipython-input-193-a12b169f134f> in dayOfYear(date)
      1 def dayOfYear(date):
      2     # YOUR CODE HERE
----> 3     year=int(re.findall('(\d+)[a-]',date)[0])
      4     month=int(re.findall('[a-](\d+)[a-]',date)[0])
      5     day=int(re.findall('[a-](\d+)',date)[-1])

IndexError: list index out of range
```

```python

```