Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Problem 1 of quiz 2 (not counted towards the final grade), STATS102 Spring 2021 Session 3, Duke Kunshan University,  Jan 13th

![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Test whether the elements in input the list are in a strictly ascending order, return bool type decision.

<!-- grade_id: strictly_ascending_list, points: 0, grade: false, solution: true -->
```python
def strictly_ascending_list(mylist):
    # YOUR CODE HERE
    if mylist
    elif mylist==sorted(mylist):
        return(True)
    else:
        return(False)
```

<!-- grade_id: strictly_ascending_list, points: 0, grade: false, solution: true -->
```python
def strictly_ascending_list(mylist):
    # YOUR CODE HERE
    if(len(mylist)==1):
        return True

    for i in range(1,len(mylist)):
        if(mylist[i])<=(mylist[i-1]):
            return False

    return True
```

```python
range(1,9)
```

Output:
```
range(1, 9)
```

<!-- grade_id: strictly_ascending_list1, points: 0.5, grade: true, solution: false -->
```python

assert strictly_ascending_list([3, 6, 8, 9])==True
assert strictly_ascending_list([-1, 3, 11, 17, 19])==True
assert strictly_ascending_list([-1, 0, 1, 2])==True
assert strictly_ascending_list([9, 8, 7, 6])==False
assert strictly_ascending_list([-1, -2, -3, -4])==False
assert strictly_ascending_list([1, 1, 4, 5])==False

```