Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Problem 1 of quiz 1 (free exercise, not counted towards to the final grade), STATS102 Fall 2022 Session 1, Duke Kunshan University, Aug 25th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

<!-- grade_id: calculate_myfactorial, points: 0, grade: false, solution: true -->
```python
def myfactorial(num):
    '''Calculate the Factorial of the input num
    the output is an int, the check part is already provided to you, please just focus on the calculation of the factorial
    '''
    if(type(num)!=int):
        raise VauleError('input is not an integer')
    if(num<=0):
        raise VauleError('input is not an integer')
    '''start your code here, num is a positive integer here'''
    # YOUR CODE HERE
    for i in range(1,num):
        if num=1
    '''end your code'''
```

<!-- grade_id: calculate_myfactorial, points: 0, grade: false, solution: true -->
```python
def myfactorial(num):
    '''Calculate the Factorial of the input num
    the output is an int, the check part is already provided to you, please just focus on the calculation of the factorial
    '''
    if(type(num)!=int):
        raise VauleError('input is not an integer')
    if(num<=0):
        raise VauleError('input is not an integer')
    '''start your code here, num is a positive integer here'''
    # YOUR CODE HERE
    fac=1
    while num>1:
        fac=fac*num
        num=num-1
        
    '''end your code'''
```

```python
range(1,9)
```

Output:
```
range(1, 9)
```

<!-- grade_id: myfactorial_correct_input, points: 0.5, grade: true, solution: false -->
```python
"""Check that calculate_income function returns the correct output for several inputs"""
assert myfactorial(7) == 5040
assert myfactorial(9) == 362880
```