Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 2, STATS102 2022 Fall Session 1, Duke Kunshan University, Due at Aug 26th 1:15PM beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

# ---
## Part A (1 points)

Write a function that returns a list of numbers, such that $x_i=(i-2)^2-4$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ or $n$ is not an integer by raising a `ValueError`.

<!-- grade_id: cubics, points: 0, grade: false, solution: true -->
```python
def parabola(n):
    """Compute each Xi from 1 to n, such that the ith element of the returned list equals (i-2)^2-4.
    """
    # YOUR CODE HERE
    if type(n)!=int:
        raise ValueError
    elif n<1:
        raise ValueError
    else:
        results = []
        for value in range(1,n + 1):
            result = (value - 2)**2-4
            results.append(result)
        return results
```

Your function should print `[-3, -4, -3, 0, 5, 12, 21, 32, 45, 60]` for $n=10$. Check that it does:

```python
parabola(10)
```

Output:
```
[-3, -4, -3, 0, 5, 12, 21, 32, 45, 60]
```

<!-- grade_id: correct_cubics, points: 0.75, grade: true, solution: false -->
```python
"""Check that cubics returns the correct output for several inputs"""
assert parabola(2) == [-3, -4]
assert parabola(6) == [-3, -4, -3, 0, 5, 12]
assert parabola(8) == [-3, -4, -3, 0, 5, 12, 21, 32]
assert parabola(11) == [-3, -4, -3, 0, 5, 12, 21, 32, 45, 60, 77]
```

<!-- grade_id: cubics_invalid_input, points: 0.25, grade: true, solution: false -->
```python
"""Check that cubics raises an error for invalid inputs"""
try:
    parabola(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    parabola(-10)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```

---

## Part B (1 point)

Using your `parabola` in Part A, write another function that computes the difference between the maximum and minimum of the returned list of `parabola`. Make sure your function call the `parabola` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_cubics, points: 0, grade: false, solution: true -->
```python
def diff_of_max_min(n):
    """Compute the difference between the maximum and minimum in the list."""
    # YOUR CODE HERE
    a=parabola(n)
    maxx=-2147483648
    minn=2147483648
    for i in range(len(a)):
        if a[i]>maxx:
            maxx=a[i]
        if a[i]<minn:
            minn=a[i]
    return maxx-minn
```

The difference between the maximum and minimum when $n$ is 10 should be 64. Verify that this is the answer you get:

```python
diff_of_max_min(10)
```

Output:
```
64
```

<!-- grade_id: correct_sum_of_cubics, points: 0.5, grade: true, solution: false -->
```python
"""Check that diff_of_max_min returns the correct answer for various inputs."""
assert diff_of_max_min(1) == 0
assert diff_of_max_min(2) == 1
assert diff_of_max_min(8) == 36
assert diff_of_max_min(11) == 81
```

<!-- grade_id: sum_of_cubics_uses_cubics, points: 0.5, grade: true, solution: false -->
```python
"""Check that diff_of_max_min relies on parabola."""
orig_parabola = parabola
del parabola
try:
    diff_of_max_min(1)
except NameError:
    pass
else:
    raise AssertionError("diff_of_max_min does not use parabola")
finally:
    parabola = orig_parabola
```