Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

##### quiz 5, STATS102 Fall 2022 Session 1, Duke Kunshan University, Aug 31th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

# Problem 1
For an input numpy array x, return a string in the following format:

'the input numpy array has xxx rows, xxx columns!'

xxx should be replaced by the real integer number.

For example:


the input 

array([[0, 1, 2],
       [3, 4, 5],
       [6, 7, 8]])
       
the output

'the input numpy array has 3 rows, 3 columns!'

<!-- grade_id: strictly_ascending_list, points: 0, grade: false, solution: true -->
```python
import numpy as np
def tell_the_shape(x):
    # YOUR CODE HERE
    a=x.shape
    row=int(a[0])
    column=int(a[1])
    return 'the input numpy array has '+str(row)+' rows, '+str(column)+' columns!'

#{'the input numpy array has %d rows, %d columns!'}.format()
```

```python
data=np.arange(12).reshape(3,4)
tell_the_shape(data)
```

Output:
```
'the input numpy array has 3 rows, 4 columns!'
```

<!-- grade_id: strictly_ascending_list1, points: 0.5, grade: true, solution: false -->
```python
import numpy as np
data=np.arange(12).reshape(3,4)
assert tell_the_shape(data)=='the input numpy array has 3 rows, 4 columns!'
data=np.arange(15).reshape(3,5)
assert tell_the_shape(data)=='the input numpy array has 3 rows, 5 columns!'
data=np.empty((4,4))
assert tell_the_shape(data)=='the input numpy array has 4 rows, 4 columns!'
```

# Problem2
for an input scalar variable x (\\( x \in (-1,1)\\)) and positive integer N, calculate and return \\( \Sigma_{n=0}^{N} (-x)^{n}\\), we will compare the results with \\( \frac{1}{1+x}\\) according to different N. Hint: np.arange np.sum

<!-- grade_id: cell-93b6d6060132db94, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myfunction(x,N):
    # YOUR CODE HERE
    n=np.arange(N+1)
    data1=(-x)**n
    return data1.sum()
```

<!-- grade_id: cell-57a0cf34b31278ec, points: 0.5, grade: true, solution: false -->
```python
import numpy as np
assert(abs(myfunction(0.5,10)-1/(1+0.5))>1e-5)
assert(abs(myfunction(0.5,50)-1/(1+0.5))<1e-5)
assert(abs(myfunction(0.5,100)-1/(1+0.5))<1e-10)
assert(abs(myfunction(0.4,5)-1/(1+0.4))>1e-5)
assert(abs(myfunction(0.4,20)-1/(1+0.4))<1e-5)
```