Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

quiz3 (not counted towards the final grade), STATS102 2022 Fall session 1, Duke Kunshan University, Aug 29th 2022
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

## Problem1

One apple is 5 dollars, one pear is 3 dollars, and 3 bananas are 1 dollar. How to buy 100 fruits using 100 dollars?

Input: no input

Output: all possible combonation of `[#apples, #pears, #bananas]`, the output should be a list of list, e.g. `[[#apples, #pears, #bananas], [#apples, #pears, #bananas], [#apples, #pears, #bananas]]`


<!-- grade_id: cell-f95bfce2ff17397a, points: 0, grade: false, solution: true -->
```python
def fun():
    for x in range(101):
        for y in range (101):
            for z in range (101):
                pay=x*5+y*3+z*(1/3)
                alist=[]
                if x+y+z==100 and pay==100:
                    aresult=[x,'apples',',',y,'pears',',',z,'bananas']
                    alist.append(aresult)
                else:
                    continue
    return alist
```

```python
fun()
```

<!-- grade_id: cell-770318cc3f2e6430, points: 0.5, grade: true, solution: false -->
```python
def my_test(results):
    for num_fruits in results:
        if len(num_fruits) != 3 or sum([a*b for a, b in zip(num_fruits, [5, 3, 1/3])]) != 100 or sum(num_fruits) != 100:
            return False
    return True
        
assert my_test(fun()) == True
```