Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Quiz 10, STATS102 Fall 2022 Session 1, Duke Kunshan University, Sep 8th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

load the census data from http://census.ire.org/data/bulkdata.html (census.csv), calcualte and return a list of top 5 state names ordered in descending way by the average county population in that state. The populations for all counties in each state are listed in column CENSUS2010POP.

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    df = pd.read_csv('census.csv')
    # YOUR CODE HERE
    df = pd.read_csv('census.csv')
    return df.groupby('STNAME').mean().sort_values(by=['CENSUS2010POP'],ascending=False).head(5).index
```

```python
import numpy as np
import pandas as pd
df = pd.read_csv('census.csv')
df.groupby('STNAME').mean().sort_values(by=['CENSUS2010POP'],ascending=False).head(5).index
```

Output:
```
Index(['California', 'District of Columbia', 'Massachusetts', 'Connecticut',
       'Arizona'],
      dtype='object', name='STNAME')
```

```python
import numpy as np
import pandas as pd
df = pd.read_csv('census.csv')
df.groupby('STNAME')['CENSUS2010POP'].mean().sort_values(ascending=False).head(5).index
```

Output:
```
Index(['California', 'District of Columbia', 'Massachusetts', 'Connecticut',
       'Arizona'],
      dtype='object', name='STNAME')
```

```python
import numpy as np
import pandas as pd
df = pd.read_csv('census.csv')
df
```

Output:
```
SUMLEV  REGION  DIVISION  STATE  COUNTY   STNAME            CTYNAME  \
0         50       3         6      1       1  Alabama     Autauga County   
1         50       3         6      1       3  Alabama     Baldwin County   
2         50       3         6      1       5  Alabama     Barbour County   
3         50       3         6      1       7  Alabama        Bibb County   
4         50       3         6      1       9  Alabama      Blount County   
...      ...     ...       ...    ...     ...      ...                ...   
3137      50       4         8     56      37  Wyoming  Sweetwater County   
3138      50       4         8     56      39  Wyoming       Teton County   
3139      50       4         8     56      41  Wyoming       Uinta County   
3140      50       4         8     56      43  Wyoming    Washakie County   
3141      50       4         8     56      45  Wyoming      Weston County   

      CENSUS2010POP  ESTIMATESBASE2010  POPESTIMATE2010  ...  \
0             545
```

<!-- grade_id: myfunction_test1, points: 0.5, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
names=myfunction()
assert names[0]=='California'
assert names[1]=='District of Columbia'
assert names[2]=='Massachusetts'
assert names[3]=='Connecticut'
assert names[4]=='Arizona'
```