Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Assignment 7, STATS102 Fall 2022 Session 1, Duke Kunshan University, Due at Sep 6th 1:15pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

## Problem1

Given an input numpy array (2D matrix), perform mean substraction for each column and return the resulted one using broadcasting

<!-- grade_id: 111, points: 0, grade: false, solution: true -->
```python
import numpy as np
def mean_sub_col(A):
    # YOUR CODE HERE
    col_mean=A.mean(axis=0)
    result=x-col_mean
    return result
```

<!-- grade_id: cell-414a358ab53a56e9, points: 1, grade: true, solution: false -->
```python
import numpy as np
x=np.arange(12).reshape(3,4)
assert (mean_sub_col(x)==np.array([[-4., -4., -4., -4.],
       [ 0.,  0.,  0.,  0.],
       [ 4.,  4.,  4.,  4.]])).all()
```

## Problem3
for an input scalar variable x (\\( x \in (0,2\pi)\\)) and positive integer N, calculate and return \\( \Sigma_{k=1}^{N} \frac{(-1)^{k-1}x^{2k-1}}{(2k-1)!}\\), we will compare the results with \\( sin(x) \\) according to different N. Hint: np.arange np.sum np.multiply.accumulate

<!-- grade_id: cell-165f4c3c5309d68b, points: 0, grade: false, solution: true -->
```python
import numpy as np
def mysin(x,N):
    # YOUR CODE HERE
    k=np.arange(1,N+1)
    n=0
    for k in range (1,N+1):
        m=((-1)**(k-1))*(x**(2*k-1))/np.math.factorial(2*k-1)
        n=n+m
    return n
```

<!-- grade_id: cell-f8edd21d15182d29, points: 1, grade: true, solution: false -->
```python
assert(abs(mysin(np.pi/2,10)-1)<1e-5)
assert(abs(mysin(np.pi/4,10)-np.sqrt(2)/2)<1e-5)
assert(abs(mysin(np.pi*3/2,10)-(-1))<1e-5)
```