Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 5, STATS102 Fall 2022 Session 1, Duke Kunshan University, Due at Sep 1st 1:15pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

# Problem1 (0.6 point)
Given a string containing just the characters `'('`, `')'`, `'{'`, `'}'`, `'['` and `']'`, determine whether the input string is complete parenthesis combination. (Empty string is also valid.)

Some correct cases:
`{[()]}`, `[]{}()`, `[()]`, `{}` ...

<!-- grade_id: cell-f55edec9336761f9, points: 0, grade: false, solution: true -->
```python
def isValid(s):
    while '{}'in s or '[]'in s or '()' in s:
        s=s.replace('[]','')
        s=s.replace('()','')
        s=s.replace('{}','')
    return s==''
```

<!-- grade_id: cell-1eddbf23c0d24eb4, points: 0.6, grade: true, solution: false -->
```python
"""Check that isValid function returns the correct output"""
assert (isValid("[[[]]{}}") == False)
assert (isValid("[](){}{[]}()[[]]") == True)
assert (isValid("[{}]") == True)
assert (isValid("([)]") == False)
assert (isValid("((((())]") == False)
```

# Problem2 (0.6 point)

Given an array of integers where 1 ≤ a[i] ≤ n (n = size of array), some elements may appear twice. Find all the elements of [1, n] inclusive that do not appear in this array. (Return `[]`, when n < 1.)

<!-- grade_id: cell-6e14d84ed557fa4d, points: 0, grade: false, solution: true -->
```python
def findDisappearedNumbers(nums):
    # YOUR CODE HERE
    n=len(nums)
    alist=[]
    if n<1:
        return []
    else:
        for i in range(1,n+1):
            if nums.count(i)==0:
                alist.append(i)
            else:
                continue
    return alist
```

<!-- grade_id: cell-4ccb27702e3e4b9d, points: 0.6, grade: true, solution: false -->
```python
assert findDisappearedNumbers([4,3,2,7,8,2,3,1,5]) == [6, 9]
assert findDisappearedNumbers([10,9,4,7,8,8,8,1,2,1,1,1,1]) == [3, 5, 6, 11, 12, 13]
assert findDisappearedNumbers([6,7,8,1,2,3,3,3,3,10]) == [4, 5, 9]
assert findDisappearedNumbers([]) == []
```

# Problem 3 (0.8 point)

Write a program to encode a string(only containing 26 letters) using symbols of international Morse code.

Reference: https://en.wikipedia.org/wiki/Morse_code

For your convienice, you can use the morse code of 26 letters below:

[".-","-...","-.-.","-..",".","..-.","--.","....","..",".---","-.-",".-..","--","-.","---",".--.","--.-",".-.","...","-","..-","...-",".--","-..-","-.--","--.."]

<!-- grade_id: cell-c033d11914475268, points: 0, grade: false, solution: true -->
```python
def encodeString(ori_str):
    # YOUR CODE HERE
    alphabet=['A','B','C','D','E','F','G','H','I','J','K','L','M','N','O','P','Q','R','S','T','U','V','W','X','Y','Z']
    mos=[".-","-...","-.-.","-..",".","..-.","--.","....","..",".---","-.-",".-..","--","-.","---",".--.","--.-",".-.","...","-","..-","...-",".--","-..-","-.--","--.."]
    d=dict(zip(alphabet,mos))
    ans=''
    for i in ori_str:
        morse=d[i]
        ans=ans+morse
    return ans
```

<!-- grade_id: cell-72c5997eaf98259b, points: 0.8, grade: true, solution: false -->
```python
assert encodeString('CUL') == '-.-...-.-..'
assert encodeString('GA') == '--..-'
assert encodeString('SIG') == '.....--.'
assert encodeString('TU') == '-..-'
```