Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Assignment 10, STATS102 Fall 2022 Session 1, Duke Kunshan University, Due at Sep 10th 1:15pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

## Problem1
Load the file (WorldCups.csv), and then return the following DataFrame sorted by the numbers in the 'Score' column. For each world cup, the winner gets 4 points, the runners-up gets 2 points, the third place gets 1 point.  ('Germany FR' needs to be converted to 'Germany' and count together)

![](assignment9.png)

<!-- grade_id: mysin, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
def myfunction():
    # YOUR CODE HERE
    df = pd.read_csv('WorldCups.csv')
    df=df.replace('Germany FR','Germany')
    df1=df['Winner'].value_counts()
    df2=df['Runners-Up'].value_counts()
    df3=df['Third'].value_counts()
    DF=pd.DataFrame({'Winner': df1,'Runners-Up': df2,'Third':df3})
    DF['Score']=DF['Winner'].replace(np.nan,0)*4+DF['Runners-Up'].replace(np.nan,0)*2+DF['Third'].replace(np.nan,0)
    return DF.sort_values(by=['Score'],ascending=False)
```

<!-- grade_id: cell-044b32eea7f84461, points: 2, grade: true, solution: false -->
```python
import numpy as np
import pandas as pd
res=myfunction()
assert(res.shape==(19, 4))
assert(res['Score'].values==np.array([28, 26, 21, 14,  8,  8,  7,  4,  4,  4,  4,  4,  2,  1,  1,  1,  1,
        1,  1])).all()

```