Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

quiz 12, STATS102 Fall 2022 Session 1, Duke Kunshan University, Sep 13th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Load <b>artists.csv</b>, calculate and return how many artists were born in the 19th century [1801-1900].

Data reference: https://www.kaggle.com/ikarus777/best-artworks-of-all-time

Hint： .str.extract('regular expression patterns')

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    # YOUR CODE HERE
    df=pd.read_csv('artists.csv')
    df=df['years'].str.extract('([a-z0-9][a-z0-9][a-z0-9][a-z0-9])\s+')
    count=0
    for i in range(49):
        if 1801<=int(df[0][i])<=1900:
            count=count+1
        else:
            continue
    return count
```

```python
import numpy as np
import pandas as pd

df=pd.read_csv('Puerto Rico Given Names.txt')
df=df[(df['Year']<=2018)&(df['Year']>=2010)]
df=df.groupby(by=['Year','Name']).sum()
df
```

Output:
```
pandas/_libs/parsers.pyx in pandas._libs.parsers.TextReader.__cinit__()

pandas/_libs/parsers.pyx in pandas._libs.parsers.TextReader._setup_parser_source()

FileNotFoundError: [Errno 2] File b'Puerto Rico Given Names.txt' does not exist: b'Puerto Rico Given Names.txt'
```

<!-- grade_id: myfunction_test1, points: 0.5, grade: true, solution: false -->
```python
assert myfunction() == 27
```