Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Problem 1 of quiz 4, STATS102 Fall 2022 Session 1, Duke Kunshan University, Aug 30th (not counted towards the final grade)
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Write a function to perform regular expression on the input string and return the converted float number that denotes the number right after the $ sign

<!-- grade_id: strictly_ascending_list, points: 0, grade: false, solution: true -->
```python
import re
def myfunction(input_string):
    # YOUR CODE HERE
    a=re.findall('\$(-*\d+\.*\d*)\D',input_string)
    return float(a[0])
```

<!-- grade_id: strictly_ascending_list1, points: 0.5, grade: true, solution: false -->
```python
a='./FER2013Test/fer0032220.png $7 0 [-0.23853336, -1.2332222, -1.2634475]'
b='./FER2013Test/fer0032223.png 0 0 [0.46080494, $-0.63641989, -0.29608598]'
c='./FER2013Test/fer0032223.png 0 $0 [0.46080494, -0.63641989, -0.29608598]'
d='./FER2013Test/fer0032222.png 3 0 [$-0.56577992, -1.3372835, 0.04990942]'
e='./FER2013Test/fer0032227.png 1 1 [0.0084493794, 0.36375782, $1.2771677]'

assert myfunction(a)==7.0
assert myfunction(b)==-0.63641989
assert myfunction(c)==0.0
assert myfunction(d)==-0.56577992
assert myfunction(e)==1.2771677
```