Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

quiz 6, STATS102 Fall 2022 Session 1, Duke Kunshan University, Sep 1st
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Given a numpy array with dtype as int, calculate how many even numbers between 5 and 15 are there. 

for example:

input:
array([[ 2,  5,  1,  4,  9],
       [ 5,  2,  1, 17, 14],
       [ 9, 16, 18,  9, 11],
       [13, 10, 17, 18,  6],
       [ 7, 11,  7,  1,  4]])


output: 3

Hint: you could solve this problem with one line code.

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myfunction(A):
    # YOUR CODE HERE
    return np.sum((A>5) & (A<15) &(A%2==0))
```

<!-- grade_id: myfunction_test, points: 0.5, grade: true, solution: false -->
```python
import numpy as np
rng = np.random.RandomState(12345)
assert myfunction(rng.randint(0,20,(5,5)))==3
rng = np.random.RandomState(12345)
assert myfunction(rng.randint(0,20,(1,5)))==0
rng = np.random.RandomState(12345)
assert myfunction(rng.randint(0,20,(5,1)))==0
rng = np.random.RandomState(12345)
assert myfunction(rng.randint(0,20,(6,3)))==2
rng = np.random.RandomState(12345)
assert myfunction(rng.randint(0,20,(2,4)))==0
rng = np.random.RandomState(12345)
assert myfunction(rng.randint(0,20,(10,10)))==25

```