Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ziyue Yin"
COLLABORATORS = ""
```

---

Problem 1 of Assignment 4, STATS102 Fall Session 1, Duke Kunshan University, Due at Aug 31th 1:15pm beijing time
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

## Part A (1.5 points)

The criminal investigation brigade are now finding criminal for a criminal case. There are 8 suspects and here are the preliminary conclusions:
1. There is at least one criminal in A and B;
2. There are at least two criminals among A, C, D and E;
3. B and E may be both criminals or they are both innocent;
4. There are at least three criminals among C, D, F and G;
5. There is only one criminals in G and H;
6. If E is innocent, H is also innocent.

Write a program to find out there are how many possible combinations of eight suspects' positions.

Hint: each combination can be expressed by a list like `[1,0,0,0,1,1,1,0]`, where the \\(i^{th}\\) element of the list is 0 if the \\(i^{th}\\) suspect is not a criminal, and the \\(i^{th}\\) element of this list is 1 if the \\(i^{th}\\) suspect is a criminal.

<!-- grade_id: sort_dict, points: 0, grade: false, solution: true -->
```python
def whoIsCriminals():
    # YOUR CODE HERE
    n=0
    for A in range(2):
        for B in range(2):
            for C in range(2):
                for D in range(2):
                    for E in range(2):
                        for F in range(2):
                            for G in range(2):
                                for H in range(2):
                                    if ((A+B==1 or A+B==2) and (A+C+D+E==2 or A+C+D+E==3 or A+C+D+E==4) and (B+E==2 or B+E==0) and (C+D+F+G==3 or C+D+F+G==4) and (G+H==1) and ((E==0 and H==0) or (E==1 and (H==0 or H==1)))):
                                        n=n+1
                                    else:
                                        continue
    return n
```

<!-- grade_id: sort_dict_test, points: 1.5, grade: true, solution: false -->
```python
assert whoIsCriminals() == 14
```

## Part B (0.5 points)

Given a string `s` consists of alphabets and empty space characters `' '`, return the length of the penultimate word in the string.

If the penultimate word does not exist, return `-1`.

<!-- grade_id: cell-9fc0fc8ab751a5c6, points: 0, grade: false, solution: true -->
```python
import re
def myString(s):
    # YOUR CODE HERE
    t=s.split(' ')
    p=t[::-1]
    for i in p:
        if len(t)==1:
            return -1
        else:
            return len(p[1])
```

<!-- grade_id: cell-afe83db1fdefe84b, points: 0.5, grade: true, solution: false -->
```python
assert myString('Hi everyone') == 2
assert myString('I love Python') == 4
assert myString('Oops!') == -1
```