Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

quiz 7, STATS102 Fall2022 Session 1, Duke Kunshan University, Sep 5th
![](dku.jpg)
For this problem set, we'll be using the Jupyter notebook:
![](jupyter.png) 

Loading `chipotle.tsv`, calculate and return the total revenue of Chipotle Restaurant.

<!-- grade_id: myfunction, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def myfunction():
    chipo = pd.read_csv('chipotle.tsv', sep = '\t')
    # YOUR CODE HERE
    
```

```python
import numpy as np
import pandas as pd
chipo = pd.read_csv('chipotle.tsv', sep = '\t')
chipo
```

Output:
```
order_id  quantity                              item_name  \
0            1         1           Chips and Fresh Tomato Salsa   
1            1         1                                   Izze   
2            1         1                       Nantucket Nectar   
3            1         1  Chips and Tomatillo-Green Chili Salsa   
4            2         2                           Chicken Bowl   
...        ...       ...                                    ...   
4617      1833         1                          Steak Burrito   
4618      1833         1                          Steak Burrito   
4619      1834         1                     Chicken Salad Bowl   
4620      1834         1                     Chicken Salad Bowl   
4621      1834         1                     Chicken Salad Bowl   

                                     choice_description  item_price  
0                                                   NaN        2.39  
1                                          [Clementine]      
```

```python
import numpy as np
import pandas as pd
chipo = pd.read_csv('chipotle.tsv', sep = '\t')
chipo['inconme']=chipo['item_price']*chipo['quantity']
chipo

```

Output:
```
order_id  quantity                              item_name  \
0            1         1           Chips and Fresh Tomato Salsa   
1            1         1                                   Izze   
2            1         1                       Nantucket Nectar   
3            1         1  Chips and Tomatillo-Green Chili Salsa   
4            2         2                           Chicken Bowl   
...        ...       ...                                    ...   
4617      1833         1                          Steak Burrito   
4618      1833         1                          Steak Burrito   
4619      1834         1                     Chicken Salad Bowl   
4620      1834         1                     Chicken Salad Bowl   
4621      1834         1                     Chicken Salad Bowl   

                                     choice_description  item_price  inconme  
0                                                   NaN        2.39     2.39  
1                                          
```

<!-- grade_id: myfunction_test, points: 0.5, grade: true, solution: false -->
```python
assert (myfunction()-39237.02 < 0.001)
```